Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Alicia"
COLLABORATORS = ""
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
import math
#import math to be able to write negative infinity
def bruteforce_max_subarray(A):
    """
    Implements brute-force maximum subarray finding.
    
    Parameters
    ----------
    A : list
        a NON-EMPTY list of floats
    
    Returns
    -------
    tuple
        - the start index of the max subarray
        - the end index of the max subarray
        - the value of the maximum subarray
        
    """
    maxi = -math.inf
    #put the lowest possible number and will update it later
    
    for i in range(0, len(A)-1):
        #define i that will go from the first element to the second last element in the array
        sumi = 0
        #define the sum for the subarrays so far
        for j in range(i, len(A)):
            #define j that we can consider as the end of our maximum array
            sumi = A[i] + A[j]
            if sumi > maxi:
                maxi = sumi
                start = i
                end = j
    return maxi, start, end

bruteforce_max_subarray([-2, -3, 4, -1, -2, 1])
```

Output:
```
(8, 2, 2)
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. With the output from ```max_subarray(A)```, do you need to rerun ```max_subarray()``` for the new array again to find the subarray we need? Explain why, or why not, in as much detail as possible, using, if necessary, a specific example to illustrate your reasoning. The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
Yes, because a brute force algorithm will need to compare every possible pair of numbers. so it is needed to re-run the whole algorithm in order to find the maximum subarray after adding the last element. For example if the array was (A, B, C), it will have the subarray of (A,B), (A,B,C), and (B,C), if we add D to the array, it need to rerun in order to include D as part of the subarray. 

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
False, the new maximum-subarray needs to include 8 because 8 is positive and is affecting the final answer of finding the maximum subarray. If the newly added element was -8 for example then we do not need to inlcude it because either ways the final answer will remain the same.

## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(x, mx):
    """ 
    Parameters
    ----------
    x : list
        A NON-EMPTY list of numbers (e.g., the 
        array B in the first two questions above)
        
        If x has 1 element: returns the value of 
        the element regardless of the value of mx
        
    mx 
        The maximum subarray of x excluding its 
        ast element (i.e., compute the maximum 
        subarray of the input array x considering 
        only its first len(x) - 1 elements)
    
    Returns
    -------
    int
        The maximum subarray of the array x
    
    Example
    -------
    Using the array B in question 2, the result 
    of incremental_max_subarray(B, 4) is 10: 
    
        10 = 8 + 1 - 2 -1 + 4 
    
    
    """
    sumi = 0
    #setting the initial total
    maxi = float('-inf')
    #setting the initial maximum
    for i in x[::-1]:
        #set i values as each value of x from the last element to the first
        sumi = sumi + i
        #adding each element to the total
        if sumi > mx and sumi > maxi:
            #comparing the total we have to the previous maximum subarray and the inital maximum for the array
            maxi = sumi
            #setting the new maximum subarray we have equal to our current total
    return maxi
incremental_max_subarray([-2, -3, 4, -1, -2, 1, 8], 4)
```

Output:
```
10
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-4-2256e7bbc067> in <module>
      1 B = [-2, -3, 4, -1, -2, 1, 8]
      2 assert(incremental_max_subarray(B, 4) == 10)
----> 3 assert(incremental_max_subarray(B[:1], 0) == B[0])

AssertionError:
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def max_subarray(A):
    """
    Using `incremental_max_subarray` iteratively on A 
    to produce the value of the maximum subarray of A.
    
    Parameters
    ----------
    A : list
        A NON-EMPTY list of floats
        
    Returns
    -------
    float
        The sum of the maximum subarray of A 
    
    """
    sumi = 0
    #set the initial value of the total
    maxi = float('-inf')
    #set the initial value of the maximum
    for i in A[::-1]:
        #put the range for i interating backwards for list A
        sumi = sumi + i
        #adding value of i to the totol
        if sumi > maxi:
            #comparing the total to the current maximum
            maxi = sumi    
            #updating the current maximum
    return maxi
#returning the maximum


max_subarray([-3,-5,6,-1,-2,9,10])
```

Output:
```
22
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
Yes. By using this approach, we have something that is already built which is "mx" that helps us in comparing it with our new subarray and return the maximum in fewer steps.

```python

```